# Funciones

Temas:
1. Como declarar una función en Julia
2. "Duck-typing" en Julia
3. Funciones mutantes vs no mutantes en Julia
4. Algunas funciones de mayor orden

## Como declarar una función en Julia
Julia tiene varias formas de declarar funciones. La primera es utilizando las palabras clave `function` y `end`.

```julia
function saludo(nombre)
    println("Hola $nombre, ¡qué bueno verte!")
end
```

```julia
function f(x)
    x^2
end
#Nota: podríamos haber escrito return x^2, y también habría funcionado
#Las funciones deJulia retornan el resultado de la última operación que
#hicieron.
```

Podemos llamar a cualquiera de estas funciones así:

```julia
saludo("Mario")
```

```julia
f(42)
```

Alternativamente podríamos haber declarado esas funciones en una sola línea

```julia
salu2(nombre) = println("Hola $nombre, ¡qué bueno verte!")
```

```julia
f2(x) = x^2
```

```julia
salu2("María")
```

```julia
f2(42)
```

Finalmente, podríamos haberlas declarado como funciones "anónimas"

```julia
saludo3 = nombre -> println("Hola $nombre, ¡qué bueno verte!")
```

```julia
f3 = x -> x^2
```

```julia
saludo3("René")
```

```julia
f3(42)
```

## Duck-typing in Julia
*"Si hace cuak como un pato, es un pato."* <br><br>
Las funciones de Julia van a funcionar siempre que el input tenga sentido. <br><br>
Por ejemplo, saludo funciona si el argumento es un número

```julia
saludo(3.14)
```

Y `f` funcionará en una matriz. 

```julia
A = rand(3, 3)
A
```

```julia
f(A)
```

`f` también va a funcionar con un string, ya que `^` es el operador utilizado para repetir strings.

```julia
f("hola")
```

Sin embargo, `f` no va a funcionar en un vector. A diferencia de `A^2`, que está bien definido, el significado de `v^2` donde `v` es un vector no es una operación algebráica bien definida.

```julia
v = rand(3)
```

```julia
# Esta línea dará error
f(v)
```

```julia

```

Output:
```
[36m  print([io::IO], xs...)[39m

  Write to [36mio[39m (or to the default output stream [36mstdout[39m if [36mio[39m is not given) a
  canonical (un-decorated) text representation. The representation used by
  [36mprint[39m includes minimal formatting and tries to avoid Julia-specific details.

  [36mprint[39m falls back to calling [36mshow[39m, so most types should just define [36mshow[39m.
  Define [36mprint[39m if your type has a separate "plain" representation. For
  example, [36mshow[39m displays strings with quotes, and [36mprint[39m displays strings
  without quotes.

  [36mstring[39m returns the output of [36mprint[39m as a string.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡≡≡[22m

[36m  julia> print("Hello World!")[39m
[36m  Hello World![39m
[36m  julia> io = IOBuffer();[39m
[36m  [39m
[36m  julia> print(io, "Hello", ' ', :World!)[39m
[36m  [39m
[36m  julia> String(take!(io))[39m
[36m  "Hello World!"[39m

signal (15): Terminado
in expression
```

## Funciones mutantes vs. no-mutantes

A esta altura habrá notado que algunas funciones en los tutoriales anteriores tienen un signo `!`. Por convención, las funciones que terminan en un signo `!` pueden alterar el contenido de sus argumentos, y las que no lo tienen no.

Por ejemplo, miremos la diferencia entre las funciones de ordenamiento `sort` y `sort!`

```julia
v = [3, 5, 2]
```

```julia
sort(v)
```

```julia
v
```

`sort(v)` retorna un arreglo ordenado de los elementos de `v`, pero `v` no cambia. <br><br>

Por otro lado, si usamos `sort!(v)`, los contenidos de `v` son ordenados dentro del mismo array.

```julia
sort!(v)
```

```julia
v
```

Cuando defina sus propias funciones trate de usar esta convención, de manera de saber rápido si la función que está utilizando corre el riesgo de modificar vectores que ha definido.

## Algunas funciones de mayor orden

### map

`map` es una función de "mayor orden" de Julia que *toma una función* como uno de sus argumentos. `map` luego aplica esa función a todos los elementos de una estructura de datos que le pases.
Por ejemplo, ejecutar

```julia
map(f, [1, 2, 3])
```
retornará un arreglo donde `f` fue aplicada a todos los elementos de `[1, 2, 3]`
```julia
[f(1), f(2), f(3)]
```

```julia
map(f, [1, 2, 3])
```

```julia

```

Es decir, elevamos al cuadrado cada elemento de `[1, 2, 3]`.

También podríamos haberle dado a `map` una función anónima tal como

```julia
x -> x^3
```

via

```julia
map(x -> x^3, [1, 2, 3])
```

y ahora hemos elevado al cubo los elementos de `[1, 2, 3]`

### broadcast

`broadcast` es otra función de mayor orden como `map`. `broadcast` es una generaliación de `map`, por lo que puede hacer lo mismo que `map` y mucho más. La sintaxis de `broadcast` es la misma que para `map`

```julia
broadcast(f, [1, 2, 3])
```

Y nuevamente aplicamos `f` a todos los elementos de `[1, 2, 3]`

Una abreviatura muy útil de `broadcast` es colocar `.` entre una función a la que quiera hacer `broadcast` y sus argumentos. Por ejemplo,

```julia
broadcast(f, [1, 2, 3])
```
es equivalente a
```julia
f.([1, 2, 3])
```

```julia
f.([1, 2, 3])
```

Note nuevamente como esto es distinto a llamar
```julia
f([1, 2, 3])
```
Podemos elevar al cuadrado cada elemento del vector, ¡pero esto no es lo mismo que elevar al cuadrado el vector, ya que esto no está definido!

Para mostrar más claramente la diferencia, considere

```julia
f(A)
```
y
```julia
f.(A)
```
para una matriz `A`:

```julia
A = [i + 3*j for j in 0:2, i in 1:3]
```

```julia
f(A)
```

En este caso estamos haciendo
```
f(A) = A^2 = A * A
``` 

Por otro lado,

```julia
B = f.(A)
```

contiene los cuadrados de cada una de las entradas individuales de `A`.

Esta sintaxis nos permite escribir expresiones complejas de una forma mucho más natural

```julia
A .+ 2 .* f.(A) ./ A
```

en lugar de

```julia
broadcast(x -> x + 2 * f(x) / x, A)
```

y ambas expresiones van a tener la misma performance.

**Pequeño paréntesis: macros**

Los macros nos permiten insertar de forma fácil expresiones en Julia, y comienzan con el símbolo `@`. La diferencia principal con una función es que los macros insertan expresiones en tiempo de compilación, y las funciones ejecutan expresiones en tiempo de cómputo. Si quiere aprender más sobre macros y expresiones puede leer la sección https://docs.julialang.org/en/v1/manual/metaprogramming/ de documentación de Julia.

Un macro que verá muchas veces en este curso es `@.`. Este macro modifica líneas de código y las convierte en broadcast. Así, al ejemplo anterior podríamos haberlo escrito de forma incluso más compacta como


```julia
@. A + 2 * f(A) / A
```

### Ejercicios
#### 6.1 
Escriba una función `suma_uno` que sume 1 a su argumento

```julia

```

```julia
@assert suma_uno(1) == 2
```

```julia
@assert suma_uno(11) == 12
```

#### 6.2 
Use `map` o `broadcast` para incrementar cada elemento de una matriz `A` en `1` y asígnelo a una variable `A1`.

```julia

```

```julia
@assert A1 == [2 3 4; 5 6 7; 8 9 10]
```

#### 6.3 
Use la sintaxis `.` de `broadcast` para incrementar todos los valores de `A1` en `1` y guarde el resultado en `A2`

```julia

```

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
@assert A2 == [3 4 5; 6 7 8; 9 10 11]
```

#### 6.4

Cree una función `misuma(vec_1, vec_2)` que retorne la suma directa de `vec_1, vec_2`. Use para esto un loop `for` dentro de la función.
Luego cree otra función llamada `misuma!(vec_r, vec_1, vec_2)` que funcione igual, pero guarde el resultado en `vec_r`.

```julia
#= 
function misuma(vec_1, vec_2)
    *código*
end
=#
```

```julia
#= 
function misuma!(vec_r, vec_1, vec_2)
    *código*
end
=#
```

```julia
#Verificación de misuma. Si esta celda no da error, misuma funciona.
vec_1 = [1,2,3]
vec_2 = [3,2,1]

for elemento in misuma(vec_1, vec_2)
    @assert elemento == 4
end
```

```julia
#Verificación de misuma!. Si esta celda no da error, misuma! funciona.
vec_1 = [1,2,3]
vec_2 = [3,2,1]
vec_r = copy(vec_1)
misuma!(vec_r, vec_1, vec_2)
for elemento in vec_r
    @assert elemento == 4
end
```

#### 6.5

Repita el ejercicio anterior, pero usando broadcast para escribir una función más compacta

```julia
vec_1 = [1,2,3]
vec_2 = [3,2,1]

for elemento in misuma(vec_1, vec_2)
    @assert elemento == 4
end
```

```julia
vec_1 = [1,2,3]
vec_2 = [3,2,1]
vec_r = copy(vec_1)
misuma!(vec_r, vec_1, vec_2)
for elemento in vec_r
    @assert elemento == 4
end
```